# Programming Assignment 2
* CSCI-5930 ML Fall 2022
* Author: Dr. B

## Rules to work with this notebook
* You can create new cells in this notebook.
* PLEASE DO NOT DELETE ANY EXISTING CELL FROM THIS NOTEBOOK.
* In a cell, you are allowed only to write codes/snippets right after the marking like ##YOUR CODE. You are not allowed to write anywhere else in that particular cell.
* Your work for each task has to be in the specific cells with the ##YOUR CODE HERE marks, otherwise your submission won't be graded. Please do not delete the marks.
* After the ##YOUR CODE HERE marks, you may see command `raise NotImplementedError()` exception that alerts us that you did not attempt a particular task. Be sure to delete that command if you attempt to solve it.

```python
#Please make sure you are using a python3.9.x interpreter 
# and install the packages from requirements.txt (provided)
!pip install -r requirements.txt
```

Output:
```
Requirement already satisfied: appnope==0.1.3 in /opt/anaconda3/lib/python3.9/site-packages (from -r requirements.txt (line 1)) (0.1.3)
Requirement already satisfied: asttokens==2.0.8 in /opt/anaconda3/lib/python3.9/site-packages (from -r requirements.txt (line 2)) (2.0.8)
Requirement already satisfied: backcall==0.2.0 in /opt/anaconda3/lib/python3.9/site-packages (from -r requirements.txt (line 3)) (0.2.0)
Requirement already satisfied: debugpy==1.6.3 in /opt/anaconda3/lib/python3.9/site-packages (from -r requirements.txt (line 4)) (1.6.3)
Requirement already satisfied: decorator==5.1.1 in /opt/anaconda3/lib/python3.9/site-packages (from -r requirements.txt (line 5)) (5.1.1)
Requirement already satisfied: entrypoints==0.4 in /opt/anaconda3/lib/python3.9/site-packages (from -r requirements.txt (line 6)) (0.4)
Requirement already satisfied: executing==1.0.0 in /opt/anaconda3/lib/python3.9/site-packages (from -r requirements.txt (line 7)) (1.0.0)
Requirement already satisfied: ipy
```

```python
#Some basic imports. You can import any package at any cell
from sklearn.preprocessing import OneHotEncoder   #My favorite categorical to numerical feature conversion tool
import pandas as pd
import numpy as np
from sklearn.preprocessing import MinMaxScaler
from sklearn.preprocessing import StandardScaler
from sklearn.impute import SimpleImputer
from sklearn.model_selection import train_test_split
from numpy.linalg import inv
from time import perf_counter
import math
import sys
import random
```

```python
### Let's load the given dataset
# First load the dataset into pandas dataframe
full_dataset = pd.read_csv('dataset/baby-weights-dataset.csv',delimiter=',')
judge_dataset = pd.read_csv('dataset/judge-without-labels.csv',delimiter=',')
```

```python
judge_dataset
```

Output:
```
ID  SEX  MARITAL  FAGE  GAINED  VISITS  MAGE  FEDUC  MEDUC  TOTALP  \
0        1    1        1    30      70      11    26     14     16       4   
1        2    2        2    21      36      15    18      9     12       1   
2        3    2        1    22      18      10    25     12     12       5   
3        4    2        1    24      25      10    22     12     11       4   
4        5    2        1    24      38      15    26     12     12       1   
...    ...  ...      ...   ...     ...     ...   ...    ...    ...     ...   
1995  1996    2        1    21      17       9    21     13     12       1   
1996  1997    1        1    27      20      10    26     16     16       2   
1997  1998    1        1    29      25      17    27     12     12       2   
1998  1999    1        1    37      29      11    34     12     16       2   
1999  2000    1        2    30      20      13    27     12     13       2   

      ...  HEMOGLOB  HYPERCH  HYPERPR  ECLAMP  CERVIX  PINFANT PRETERM 
```

## TASK 1: 
Separate the full_dataset into two parts: X and y, where X denotes the input matrix containing only the input (i.e., independent explanatory) variables, and y denotes the target variable containing only the target values for exactly the same number of samples in the given full_dataset. 

<!-- grade_id: cell-task-1-ans, points: 0, grade: false, solution: true -->
```python
#Make sure you write your solution to this task below in this cell only.

# YOUR CODE HERE
import pandas as pd
X = full_dataset.iloc[:,:-1]
y = full_dataset.iloc[:,-1]

```

<!-- grade_id: cell-task-1-test, points: 10, grade: true, solution: false -->
```python
#Let me test your code above
assert X.shape==(101400,36)
assert y.shape==(101400,)
```

```python
y.head()
```

Output:
```
0    4.3750
1    6.9375
2    8.5000
3    8.5000
4    9.0000
Name: BWEIGHT, dtype: float64
```

## TASK 2:
* Given X representing the input matrix from the full_dataset, y being the target vector (the rightmost column of the full_dataset), obtained from Task 1: 
* randomly split the (X,y) dataset into 75% for training and 25% for testing using the library function from the library [sklearn.model_selection](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html) . Please pass to the train_test_split function an additional argument random_state=45931.
* Store the 4 splits as X_train, X_test, y_train, y_test respectively.
* Save the ID column for X_train and X_test into ID_train and ID_test as list variable.
* Now, drop the ID columns from both X_train and X_test

<!-- grade_id: cell-task-2-ans, points: 0, grade: false, solution: true -->
```python
#Make sure you write your solution to this task below in this cell only.

X_train = []
X_test = []
y_train = []
y_test = []

# YOUR CODE HERE
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size = 0.25,random_state = 45931)
ID_train = X_train[['ID']]
ID_test = X_test[['ID']]
X_train = X_train.drop(['ID'], axis=1)
X_test = X_test.drop(['ID'], axis=1)

```

<!-- grade_id: cell-task-2-test, points: 10, grade: true, solution: false -->
```python
#Let me test your code above
assert (X_train.shape, X_test.shape, y_train.shape, y_test.shape) == ((76050, 35), (25350, 35), (76050,), (25350,))
```

## TASK 3:
Compute mean, stdev, min, max, 25% percentile, median and 75% percentile of BWEIGHT target variable (i.e, the target y) in the training set (i.e., y_train), and print the computed values as a numpy array containing these 7 results (respectively).


<!-- grade_id: cell-task-3-ans, points: 0, grade: false, solution: true -->
```python
#Make sure you write your solution to this task below in this cell only.
mean_val = 0
stdev_val = 0
min_val = 0
max_val = 0
percentile25_val = 0
median_val = 0
percentile75_val = 0

# YOUR CODE HERE

mean_val = np.mean(y_train)
stdev_val = np.std(y_train)
min_val = np.min(y_train)
max_val = np.max(y_train) 
percentile25_val = y_train.quantile(0.25) 
median_val = y_train.quantile(0.5) 
percentile75_val = y_train.quantile(0.75)

print( np.array([mean_val,stdev_val, min_val,max_val,percentile25_val,median_val,percentile75_val]))
```

Output:
```
[ 7.25699786  1.33004969  0.3125     13.0625      6.625       7.375
  8.0625    ]
```

<!-- grade_id: cell-task-3-test, points: 10, grade: true, solution: false -->
```python
assert abs(sum([mean_val,stdev_val, min_val,max_val, percentile25_val,median_val,percentile75_val])-44.024547552217015) < 1e-4
```

## TASK 4: 

A little background first: Categorical features are features that contain values that are not numeric. As you can imagine non-numeric values will create trouble (by introducing `nan`) during calculation to gradients, and whatnot, right? The maths are undefined when these get in its way. An obvious solution you may be intrigued to do is dropping the features! Aha! Wrong!! Every piece of data is precious... as those non-numeric features may present with valuable insights of the data samples to find the patterns to map inputs with output/targets. So, we should include them. But, how?

The answer is via "Encoding" we can make good use the non-numeric variables.

There are several types of encoding used in practice. Here are the two popular ones:

1. **Label Encoding**, where labels are encoded as subsequent numbers. Say, for a categorical feature named "Category" with three categorical values: {“Cat”, “Dog” or “Zebra”} can be encoded to "0", "1", "2" respectively as in figure below. The issue with this type of encoding may unintentionally impose a type of ordering of the categories, that may add bias to the training.
![label-encoding](http://54.160.44.72/STATIC_FIGS_DO_NOT_MOVE/le.png)
2. **One Hot Encoding**, ignores the ordering of the categories all together. With one-hot, we convert each categorical value into a new categorical column and assign a binary value of 1 or 0 to those columns. Each integer value is represented as a binary vector. All the values are zero, and the index is marked with a 1. Also, don't forget to remove the original categorical features. Here below just an example, how to convert the categorical feature called "Category" having the {“Cat”, “Dog” or “Zebra”} values into three new binary features: "Cat", "Dog", "Zebra".
![label-encoding](http://54.160.44.72/STATIC_FIGS_DO_NOT_MOVE/ohe.png)

**A note on the Dummy Variable Trap**
The Dummy Variable Trap occurs when two or more dummy variables created by one-hot encoding are highly correlated (i.e., becomes multi-collinear). This means that one variable can be predicted from the others, making it difficult to interpret predicted coefficient variables in regression models. In other words, the individual effect of the dummy variables on the prediction model can not be interpreted well because of multicollinearity.

Using the one-hot encoding method, a new dummy variable is created for each categorical variable to represent the presence (1) or absence (0) of the categorical variable. For example, if tree species is a categorical variable made up of the values pine, or oak, then tree species can be represented as a dummy variable by converting each variable to a one-hot vector. This means that a separate column is obtained for each category, where the first column represents if the tree is pine and the second column represents if the tree is oak. Each column will contain a 0 or 1 if the tree in question is of the column's species. These two columns are multi-collinear since if a tree is pine, then we know it's not oak and vice versa. The machine learning models trained on dataset having this multi-collinearity suffers. A remedy is to drop first (or any one) of the dummy (i.e., one-hot) features created.

Given the training dataset (X_train, y_train), save as X_train_ohe after replacing all the non-numeric variables (i.e., categorical variables) with numeric encoding. Also, you need to use the same encoding to work on the X_train dataset as well and save it as X_test_oh. We need to make sure any new categorical value in the test dataset get ignored, i.e, when an unknown category is encountered during transform, the resulting one-hot encoded columns for this feature will be all zeros. Please complete the following `almost complete` function `lets_do_one_hot_encoding()` that encodes dataset on the supplied categorical feature column list.


<!-- grade_id: cell-task-4-ans-1, points: 0, grade: false, solution: true -->
```python
#Make sure you write your solution to this task below in the marked location in this cell only.


def lets_do_one_hot_encoding(data, categorical_features, transform_only=True, encoders=[], verbose=False):
    '''
    The function does one_hot_encoding on the given dataset... 
    It's intentionally defined fully, meaning you do not need to do anything here... but show us
    how to use it properly.
    
    Input: 
        * data -- it's pretty much either X_train, or X_test that you prepared previously, that is
                  any dataframe having independent variables.
        * categorical_features -- a list of column/feature names in the data (dataframe) that you think
                  are non-numerical / i.e., categorical that you want to be converted to numerical
                  using the one-hot encoding technique.
        * transform_only -- a boolean parameter, if set to False, will learn (i.e., fit) various categorical
                values in the categorical_features from the given dataset, and use this to encode the dataset
                using one-hot encoding. This is important that you set to False on training dataset, and
                True on test set. If new categorical values are present in the test dataset, those will be
                ignored, making it easier to have same set of encoded features both in training and test 
                dataset... otherwise, subsequent operations (may/) will not work.
        * encoders -- a list of one-hot encoders previously saved, and now will be used to encode given dataset.
                If transform_only=True, the function looks for this provided list of encoders to encode the 
                dataset instead of learning new categories. Once again, it's expected that you pass the set
                of encoders for each of the categorical features that you encoded the training set -- i.e.,
                leave it empty for training set, and pass the set of encoders while encoding test set.
    Returns:
        * data -- the converted dataframe after the encoding is completed.
        * enc_list -- is the list of encoders used to encode the given data.
        * categorical_features -- is the list of categorical features that you would want to encode.
        It's expected that for training dataset, you save the encoder list and the list of features for later
        use to encode a test dataset.
    '''
    if len(encoders)>0:
        enc_list=encoders
    else:
        enc_list = []
    col_onehot = []
    i = 0
    for feature in categorical_features:
        if verbose: print('Dealing with feature ={}'.format(feature))
        if transform_only==True: #for test dataset
            enc = enc_list[i]
            c_onehot = enc.transform(data[[feature]])
            i = i + 1
        else: #fit and transform
            enc = OneHotEncoder(handle_unknown='ignore', drop='first', sparse=False)
            c_onehot = enc.fit_transform(data[[feature]])
            if verbose: print('c_onehot = {}'.format(c_onehot))
            enc_list.append(enc)
        c_onehot = pd.DataFrame(c_onehot, columns=list(enc.categories_[0][1:])) #dropped first column
        if verbose: print('dropped first column...c_onehot = {}'.format(c_onehot))
        c_onehot = c_onehot.add_prefix(feature+'_')
        col_onehot.append(c_onehot)
        if verbose: print('col_onehot = {}'.format(col_onehot))
    
    #concat all onehot feature columns
    concat_df = pd.concat(col_onehot,axis=1)
    #match index with given data
    concat_df.index = data.index
    
    #Here below is your task:
    #i) drop the categorical feature columns from dataframe `data`
    #ii) and then merge with those new onehot features stored in `cocat_df`
    # YOUR CODE HERE
    for feature in categorical_features:
        data =data.drop(feature, axis='columns')
    data =pd.concat([data, concat_df], axis=1, join="inner")
   
    
    return data,enc_list,categorical_features
```

<!-- grade_id: cell-task-4-ans-2, points: 0, grade: false, solution: true -->
```python
'''
let's apply one-hot encoding on two columns: "HISPMOM" and "HISPDAD" of the training dataset, X_train.
Please call the function, lets_do_one_hot_encoding() you helped defining above passing appropriate arguments.
Also, receive the 3 return values from the function call: X_train_ohe, enc_list and categorical_features to 
use them in the next question, that is, encoding the test dataset, X_test, and 
later in Task 17 to encode the judge set. 

It's always a good idea to save the list of encoders and features in a file (e.g., joblib), if you would 
want to apply your model to evaluate/predict on a new test sample. In a summary: if your model was trained with
transformed one-hot encoded dataset, a new test data also needs to be encoded with the same list of variables
defined by the encoding.

'''
X_train_ohe = []  #one-hot encoded training dataset
enc_list = [] #encoder list to be created during one-hot encoding of the training dataset
categorical_features = ['HISPMOM', 'HISPDAD'] #The list of categorical features in question.

#Here below is your task:one-hot encoding of train set: X_train
# YOUR CODE HERE
X_train_ohe, enc_list, categorical_features= lets_do_one_hot_encoding(X_train, categorical_features, False)

```

<!-- grade_id: cell-task-4-ans-3, points: 0, grade: false, solution: true -->
```python
'''
let's apply one-hot encoding on the same two columns: "HISPMOM" and "HISPDAD" of the test dataset, X_test.
Please call the function, lets_do_one_hot_encoding() you helped defining above with appropriate parameters.
Be sure to use the encoders from the training set and turn on the `transform_only` switch to true.
Also, receive the 3 return values from the function call in X_test_ohe, enc_list, categorical_features. 
Although, the second and third returned values can be ignored.

Please make sure, the one-hot encoded training set and test set has the exact same number of features/columns,
and in the same order. If not, the test in the next cell will almost certainly fail, and you will lose points,
and the training and testing of the machine learning model will absolutely fail. So, please take close
attention.

'''
X_test_ohe = []  #one-hot encoded test dataset

#Here below is your task: one-hot encoding of test set: X_test
# YOUR CODE HERE

X_test_ohe, enc_list, categorical_features= lets_do_one_hot_encoding(X_test, categorical_features, True, enc_list)


```

<!-- grade_id: cell-task-4, points: 10, grade: true, solution: false -->
```python
assert len(X_train_ohe.columns)== len(X_test_ohe.columns) and len(X_train_ohe.columns) == sum([1 for i, j in zip(X_train_ohe.columns, X_test_ohe.columns) if i == j])
assert (X_train_ohe.shape, X_test_ohe.shape)==((76050, 45), (25350, 45))
```

## TASK 5: 
* Given the X_train_ohe (Onehot encoded Pandas Dataframe from Task 4), check if there are missing values, and if yes, count how many, and impute the missing values with corresponding mean values. 
* Finally, print the counting result as a Pandas dataframe named "missing_counts" having 2 columns {variable_name,num_of_missing_values).  Please make sure that the result lists all the input variables in the given dataset. 
* Now, impute the missing values by mean of the respective variable and save the revised dataframe as X_train_ohe_imputed.

<!-- grade_id: cell-task-5-ans, points: 0, grade: false, solution: true -->
```python
missing_counts = pd.DataFrame()
X_train_ohe_imputed = X_train_ohe.copy()


#Your task below. 
# i) Please prepare the missing_counts dataframe accordingly.
# ii) impute missing value in a variable with corresponding mean of the variable.
# YOUR CODE HERE

lis =[]
for i in X_train_ohe_imputed.columns:
    if X_train_ohe_imputed[[i]].isnull().values.any():
        mean= X_train_ohe_imputed[i].mean()
        x =[]
        x.append(i)
        ct =0
        for j in X_train_ohe_imputed.index:
            if np.isnan(X_train_ohe_imputed[i][j]):
                X_train_ohe_imputed[i][j]= mean
                ct = ct+1
        x.append(ct)
        lis.append(x)
    
missing_counts = pd.DataFrame(data = lis, columns = ['variable_name', 'num_of_missing_values'])
print(missing_counts)


```

Output:
```
/var/folders/8w/fqk30xsj2tg4p5hwsm8v1qsc0000gn/T/ipykernel_71457/1363415992.py:19: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  X_train_ohe_imputed[i][j]= mean
/var/folders/8w/fqk30xsj2tg4p5hwsm8v1qsc0000gn/T/ipykernel_71457/1363415992.py:19: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  X_train_ohe_imputed[i][j]= mean
/var/folders/8w/fqk30xsj2tg4p5hwsm8v1qsc0000gn/T/ipykernel_71457/1363415992.py:19: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-vers
```

<!-- grade_id: cell-task-5-test, points: 10, grade: true, solution: false -->
```python
assert sum(missing_counts['num_of_missing_values'])==5
assert np.isnan(X_train_ohe.loc[[1748]][['FEDUC']].iloc[0,0])==True
assert np.isnan(X_train_ohe_imputed.loc[[1748]][['FEDUC']].iloc[0,0])==False
```

## TASK 6: 
* Given a X_train_ohe_imputed (Pandas dataframe from Task 5) where all the categorical variables are already replaced with numeric values, print a list of top 20 highly correlated variables with respect to the target variable, and save the result as a Pandas dataframe named top20_df with 2 columns {variable,corr_score}. 
* Here, the corr_score between a variable x and the target variable y needs to be computed using the Pearson Correlation Coefficient (PCC). Please note, PCC ranges between -1 to +1. PCC score 0 means no correlation, while value towards +1 and -1 represent positive and negative correlations respectively. For instance, PCC=0.8 and PCC=-0.8 tell similar strength positive and negative correlations between the two subject variables.
* Please do not include BWEIGHT in the top20_df list of top 20 correlated variable list.

<!-- grade_id: cell-task-6-ans, points: 0, grade: false, solution: true -->
```python
top20_df = pd.DataFrame() #You need to save the top-20 most correlated variables with respect to BWEIGHT target

# Your task below is:
# i) on the imputed dataset you got from task 5, find the top-20 most-correlated variables
# with respect to the target variable.

# YOUR CODE HERE
Xy_train = X_train_ohe_imputed.join(y_train)
corr_matrix = Xy_train.corr(method='pearson')
# Correlation with output variable
corr_target = abs(corr_matrix[["BWEIGHT"]])
corr_target_sorted = corr_target.sort_values(by="BWEIGHT", ascending=False)
top20_variables = corr_target_sorted.index[1:21] 
top20_df = pd.DataFrame({'variable':top20_variables,'corr_score':list(corr_target_sorted.iloc[1:21,0])})
top20_df
```

Output:
```
variable  corr_score
0     WEEKS    0.565254
1    GAINED    0.174466
2    VISITS    0.129247
3   HYPERPR    0.111375
4   MARITAL    0.106297
5       SEX    0.091260
6    CIGNUM    0.086711
7   RACEDAD    0.084284
8   RACEMOM    0.078811
9   PRETERM    0.075237
10   CERVIX    0.068751
11     MAGE    0.068550
12  PINFANT    0.065963
13   ECLAMP    0.065429
14    MEDUC    0.055249
15     FAGE    0.051981
16    FEDUC    0.051537
17   HYDRAM    0.050019
18  HYPERCH    0.047094
19  UTERINE    0.040339
```

<!-- grade_id: cell-task-6-test, points: 10, grade: true, solution: false -->
```python
assert top20_df.shape==(20,2)
```

## TASK 7: 
Given the X_train_ohe_imputed (as Pandas dataframe from task 5) and and top20_df (as Pandas Dataframe from Task 6) having 2 columns {variable_name,corr_score} similar to the one you computed in Task 6:
* Please save as X_train_t20 keeping only the columns listed in the top20_df dataframe.
* Repeat the process for X_test_ohe (obtained from task 4), and save it as X_test_t20.

<!-- grade_id: cell-task-7-ans, points: 0, grade: false, solution: true -->
```python
#Your task below is to slice training and test dataset to retain only the top-20 most-correlated variables
# with respect to the target variable

X_train_t20 = pd.DataFrame()
X_test_t20 = pd.DataFrame()

# YOUR CODE HERE
X_train_t20 = X_train_ohe_imputed[top20_df.iloc[:,0]]
```

```python
X_test_t20 = X_test_ohe[top20_df.iloc[:,0]]
X_test_t20.columns
```

Output:
```
Index(['WEEKS', 'GAINED', 'VISITS', 'HYPERPR', 'MARITAL', 'SEX', 'CIGNUM',
       'RACEDAD', 'RACEMOM', 'PRETERM', 'CERVIX', 'MAGE', 'PINFANT', 'ECLAMP',
       'MEDUC', 'FAGE', 'FEDUC', 'HYDRAM', 'HYPERCH', 'UTERINE'],
      dtype='object')
```

<!-- grade_id: cell-task-7-test, points: 10, grade: true, solution: false -->
```python
assert (X_train_t20.shape,X_test_t20.shape)==((76050, 20), (25350, 20))
assert len(X_train_t20.columns)== len(X_test_t20.columns) and len(X_train_t20.columns) == sum([1 for i, j in zip(X_train_t20.columns, X_test_t20.columns) if i == j])
```

## TASK 8: 
* Apply min-max scaling on the training dataset (X_train_t20 obtained from Task 7). Save the result as X_train_scaled_mm.
* Then scale the test dataset (X_test_t20 obtained from Task 7) based on the metrics you obtain when you scale the training dataset. Save the result as X_test_scaled_mm.
* PLEASE DO NOT SCALE y_train and y_test.


<!-- grade_id: cell-task-8-ans, points: 0, grade: false, solution: true -->
```python
X_train_scaled_mm = np.array(X_train_t20)
X_test_scaled_mm = np.array(X_test_t20)

#Your task below:
# YOUR CODE HERE
scaler = MinMaxScaler()
X_train_scaled_mm = scaler.fit_transform(X_train_t20) 
X_test_scaled_mm = scaler.transform(X_test_t20)
```

<!-- grade_id: cell-task-8-test, points: 10, grade: true, solution: false -->
```python
assert abs(sum(np.array([min(X_train_scaled_mm[:,0]),max(X_train_scaled_mm[:,0]),min(X_test_scaled_mm[:,0]),max(X_test_scaled_mm[:,0])])-np.array([0.0,1.0,0.0,1.0])))<1e-4

```

## TASK 9: 
* Apply standardization (i.e., normalization) scaling on the training dataset (X_train_t20 obtained from Task 7). Save the result as X_train_scaled_std.
* Then scale the test dataset (X_test_t20 obtained from Task 7) based on the metrics you obtain when you scale the training dataset. Save the result as X_test_scaled_std.
* PLEASE DO NOT SCALE y_train and y_test.




<!-- grade_id: cell-task-9-ans, points: 0, grade: false, solution: true -->
```python
X_train_scaled_std = np.array(X_train_t20)
X_test_scaled_std = np.array(X_test_t20)

#Your task below:
# YOUR CODE HERE
scaler = StandardScaler()
X_train_scaled_std = scaler.fit_transform(X_train_t20) 
X_test_scaled_std = scaler.transform(X_test_t20)
```

<!-- grade_id: cell-task-9-test, points: 10, grade: true, solution: false -->
```python
for i in np.arange(20):
    assert abs(X_train_scaled_std[:,i].mean()-0.0)<1e-4 and abs(X_train_scaled_std[:,i].std()-1.0)<1e-4

```

## TASK 10: 
Given the (X_train_scaled_std, y_train) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **linear_regression_closed_form_training** : It fits a linear regression model using the closed-form solution to obtain the coefficients, beta's, as a numpy array of m+1 values (Please recall class lecture), where *m* is the number of variables kept in X_train (the first argument to the function). Please measure the cpu_time needed during the training step. cpu_time is not equal to the wall_time. So, use time.perf_counter() for an accurate measurement. Documentation on this function can be found here: https://docs.python.org/3/library/time.html . Finally, the function returns betas (i.e., the m+1 beta values) and the  cpu_time.
* **linear_regression_closed_form_predict**: It takes a list of m+1 beta values (i.e., betas returned from the corresponding training function, and X_test (containing test samples each having *m* input variables). Now, using the provided beta values, predict each of the test samples provided, and let's name your prediction "y_pred". Return y_pred from the function.
* **RMSLE**: It takes two lists: y_test, y_pred, where the first list represents ground truth (i.e., actual) target values for the given samples, and the second list represents a corresponding predicted values for exactly same number of samples in y_test. Compute and return the Root Mean Squared Logarithmic Error (RMSLE) of the prediction. 

* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES THE LINEAR REGRESSION or RMSLE calculation.
* Now, call linear_regression_closed_form_training() function providing X_train_scaled_std, y_train obtained from Task 9, and save the returned results as betas_closed_form,cpu_time_closed_form.
* Print betas_closed_form, cpu_time_closed_form
* Call linear_regression_closed_form_predict() function providing betas_closed_form,X_test_scaled_std obtained in Task 9. Save the returned result as y_pred.
* Call RMSLE() function providing y_test and y_pred. Save returned result as rmsle_closed_form.
* Print rmsle_closed_form.

<!-- grade_id: cell-task-10-ans-1, points: 0, grade: false, solution: true -->
```python
def linear_regression_closed_form_training(X_train, y_train):
    betas = []
    cpu_time = 0

    #Your work here
    # YOUR CODE HERE
    X = np.hstack((np.ones(len(X_train)).reshape(len(X_train),1),X_train))
    y = y_train
    t1_start = perf_counter()
    betas = np.dot(np.dot(inv(np.dot(X.T,X)),X.T),y) 
    t1_stop = perf_counter()
    cpu_time = t1_stop - t1_start
    
    return betas,cpu_time

def linear_regression_closed_form_predict(betas, X_test):
    y_pred = []
    #your work below
    # YOUR CODE HERE
    X_test_canonical = np.hstack((np.ones(len(X_test)).reshape(len(X_test),1),X_test))
    betas_reshaped = betas.reshape(X_test_canonical.shape[1],1)
    y_pred = np.dot(X_test_canonical, betas_reshaped)
    
    return y_pred

def RMSLE(y_test, y_pred, verbose=False):
    rmsle_val = 0
    #Your work below
    # YOUR CODE HERE
    #y_test = np.array(y_test).reshape(y_pred.shape)
    rmsle_val = np.sqrt(np.mean(np.square(np.log1p(y_test)-np.log1p(y_pred))))
    return rmsle_val

```

<!-- grade_id: cell-task-10-ans-2, points: 0, grade: false, solution: true -->
```python
betas_closed_form = []
cpu_time_closed_form = 0
#Your task here, call appropriate function to get these two
# YOUR CODE HERE
print(y_train)
betas_closed_form, cpu_time_closed_form = linear_regression_closed_form_training(X_train_scaled_std, y_train)
betas_closed_form, cpu_time_closed_form
```

Output:
```
16961    9.0000
28131    9.0625
35114    7.1250
85941    6.8125
26819    7.1875
          ...  
62226    8.6250
87904    8.5625
32338    8.4375
27127    8.0625
28789    8.4375
Name: BWEIGHT, Length: 76050, dtype: float64
(array([ 7.25699786,  0.71519322,  0.16451626,  0.03403515, -0.08264272,
        -0.07634788, -0.12961155, -0.09253333, -0.05847441, -0.04475425,
        -0.03880411, -0.02739149,  0.10283456,  0.0723325 , -0.03912292,
        -0.01662473,  0.00921883, -0.01519333, -0.04398384, -0.02638443,
        -0.01920356]),
 0.008567458999991118)
```

<!-- grade_id: cell-task-10-ans-3, points: 0, grade: false, solution: true -->
```python
y_pred = []
#Your task here. call appropriate function to get get y prediction for the supplied test dataset scaled
# YOUR CODE HERE
y_pred = linear_regression_closed_form_predict(betas_closed_form,X_test_scaled_std)

```

<!-- grade_id: cell-task-10-ans-4, points: 0, grade: false, solution: true -->
```python
rmsle_closed_form = 0
#Your task below. Call the function to compute RMSLE score of the y predictions
# YOUR CODE HERE
#y_pred
rmsle_closed_form = RMSLE(y_test, y_pred.flatten())
```

<!-- grade_id: cell-task-10-test, points: 10, grade: true, solution: false -->
```python
assert abs(RMSLE([1,2,3,4,5],[1,1,1,1,1])-0.733674673672524)<1e-4
assert abs(RMSLE([100, 200, 300, 400, 500], [90, 190, 290, 390, 490])-0.05596497907273607)<1e-4
assert (betas_closed_form.shape[0],y_pred.shape[0])==(21,25350)
```

## TASK 11: 
Given the (X_train_scaled_std, y_train) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **linear_regression_gd_batch_training** : It fits a linear regression model using the batch gradient descent algorithm to obtain the coefficients, beta's, as a numpy array of m+1 values, where *m* is the number of variables kept in X_train (the first argument to the function). Make sure you compute average of gradients in the batch. Please use the alpha (i.e, the learning rate) and nEpoch (number of epochs) parameters in your implementation of the gradient descent algorithm. Please measure the cpu_time needed during the training step. cpu_time is not equal to the wall_time. So, use time.perf_counter() for an accurate measurement. Documentation on this function can be found here: https://docs.python.org/3/library/time.html . Finally, the function returns betas (i.e., the m+1 beta values) and the  cpu_time.
* **linear_regression_gd_batch_predict**: It takes a list of m+1 beta values (i.e., betas returned from the corresponding training function, and X_test (containing test samples each having *m* input variables). Now, using the provided beta values, predict each of the test samples provided, and let's name your prediction "y_pred". Return y_pred from the function. 

* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES THE LINEAR REGRESSION.
* Now, call linear_regression_gd_batch_training() function providing X_train_scaled_std, y_train obtained from Task 9, and alpha=0.01,nEpoch=1000, and save the returned results as betas_batch,cpu_time_batch.
* Print betas_batch, cpu_time_batch
* Call linear_regression_gd_batch_predict() function providing betas_batch,X_test_scaled_std obtained in Task 9. Save the returned result as y_pred.
* Call RMSLE() function providing y_test and y_pred. Save returned result as rmsle_batch.
* Print rmsle_batch.



<!-- grade_id: cell-task-11-ans-1, points: 0, grade: false, solution: true -->
```python
def linear_regression_gd_batch_training(X_train,y_train, alpha, nEpoch):
    random.seed(554433)
    betas = []
    cpu_time = 0
    epsilon = 1e-6
    #Your task below
    # YOUR CODE HERE
    y_pred = []
    y = np.array(y_train).reshape((len(y_train),1))
    X = np.array(X_train)
    if all(X[:,0]==1) == False:
        X = np.hstack((np.ones(len(X)).reshape(len(X), 1), X))
    betas = np.random.rand(X.shape[1]).reshape((X.shape[1],1))
    convergence_reached = False
    counter = 0
    betas_prev = betas
    
    t1_start = perf_counter()
    while not convergence_reached:
        #Compute grad
        error = np.dot(X, betas) - y
        grad = (1/len(X))*np.dot(X.T, error)
        betas = betas - alpha * grad
        if counter > nEpoch:
            convergence_reached = True
        elif np.sum(np.abs(betas-betas_prev)) < epsilon:
            convergence_reached = True
        betas_prev = betas
        counter = counter + 1
    t1_stop = perf_counter()
    cpu_time = t1_stop - t1_start    

    return betas, cpu_time

def linear_regression_gd_batch_predict(betas,X_test):
    y_pred = []
    
    #Your task below
    # YOUR CODE HERE
    X_test_canonical = np.hstack((np.ones(len(X_test)).reshape(len(X_test),1),X_test))
    betas_reshaped = betas.reshape(X_test_canonical.shape[1],1)
    y_pred = np.dot(X_test_canonical, betas_reshaped)
    
    return y_pred
```

<!-- grade_id: cell-task-11-ans-2, points: 0, grade: false, solution: true -->
```python
betas_batch = []
cpu_time_batch = 0

#your task below
# YOUR CODE HERE
betas_batch,cpu_time_batch =linear_regression_gd_batch_training(X_train_scaled_std, y_train,alpha=0.01,nEpoch=1000)
betas_batch, cpu_time_batch
```

Output:
```
(array([[ 7.25670258],
        [ 0.71507003],
        [ 0.16382044],
        [ 0.03411903],
        [-0.08254211],
        [-0.07633532],
        [-0.12969601],
        [-0.09306953],
        [-0.0622522 ],
        [-0.04157257],
        [-0.0387495 ],
        [-0.02735647],
        [ 0.08627115],
        [ 0.07252266],
        [-0.03919517],
        [-0.00929934],
        [ 0.02460225],
        [-0.02011973],
        [-0.04410028],
        [-0.02625073],
        [-0.01915461]]),
 0.6940950420000007)
```

<!-- grade_id: cell-task-11-ans-3, points: 0, grade: false, solution: true -->
```python
y_pred = []

#your task below
# YOUR CODE HERE
y_pred = linear_regression_gd_batch_predict(betas_batch,X_test_scaled_std)
```

<!-- grade_id: cell-task-11-ans-4, points: 0, grade: false, solution: true -->
```python
rmsle_batch = 0
# YOUR CODE HERE
#print(y_test)
rmsle_batch = RMSLE(y_test,y_pred.flatten())
print(rmsle_batch)

```

Output:
```
0.1390895874686932
```

<!-- grade_id: cell-task-11-test, points: 10, grade: true, solution: false -->
```python

```

## TASK 12:
Given the (X_train_scaled_std, y_train) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **linear_regression_gd_stochastic_training** : It fits a linear regression model using the stochastic gradient descent algorithm to obtain the coefficients, beta's, as a numpy array of m+1 values, where *m* is the number of variables kept in X_train (the first argument to the function). Please use the alpha (i.e, the learning rate), nEpoch (number of epochs), nIteration (number of iterations) parameters in your implementation of the gradient descent algorithm. Please measure the cpu_time needed during the training step. cpu_time is not equal to the wall_time. So, use time.perf_counter() for an accurate measurement. Documentation on this function can be found here: https://docs.python.org/3/library/time.html . Finally, the function returns betas (i.e., the m+1 beta values) and the  cpu_time.
* **linear_regression_gd_stochastic_predict**: It takes a list of m+1 beta values (i.e., betas returned from the corresponding training function, and X_test (containing test samples each having *m* input variables). Now, using the provided beta values, predict each of the test samples provided, and let's name your prediction "y_pred". Return y_pred from the function. 

* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES THE LINEAR REGRESSION.
* Now, call linear_regression_gd_stochastic_training() function providing X_train_scaled_std, y_train obtained from Task 9, and alpha=0.001,nEpoch=100, nIteration=10000 , and save the returned results as betas_stochastic,cpu_time_stochastic.
* Print betas_stochastic, cpu_time_stochastic
* Call linear_regression_gd_stochastic_predict() function providing betas_stochastic,X_test_scaled_std obtained in Task 9. Save the returned result as y_pred.
* Call RMSLE() function providing y_test and y_pred. Save returned result as rmsle_stochastic.
* Print rmsle_stochastic.


<!-- grade_id: cell-task-12-ans-1, points: 0, grade: false, solution: true -->
```python
def linear_regression_gd_stochastic_training(X_train,y_train,alpha,nEpoch,nIteration):
    random.seed(554433)
    betas = []
    cpu_time = 0
    epsilon = 1e-6

    ## YOUR CODE HERE
    # YOUR CODE HERE
    y = np.array(y_train).reshape((len(y_train),1))
    X = np.array(X_train)
    if all(X[:,0]==1) == False:
        X = np.hstack((np.ones(len(X)).reshape(len(X), 1), X))
    betas = np.random.rand(X.shape[1]).reshape((X.shape[1],1))
    convergence_reached = False
    betas_prev = betas
    t1_start = perf_counter()
    counter = 0
    for epoch in np.arange(nEpoch):
        np.random.shuffle(X)
        for i in np.arange(len(X)):
            Xi = X[i,:].reshape((1,len(X[i,:])))
            error = np.dot(Xi,betas)-y[i]
            grad = np.dot(Xi.T,error)
            betas = betas - alpha * grad
            if counter > nIteration:
                convergence_reached = True
            elif np.sum(np.abs(betas-betas_prev)) < epsilon:
                convergence_reached = True
            betas_prev = betas
            counter = counter + 1
            if convergence_reached:
                break
        if convergence_reached:
            break
    t1_stop = perf_counter()
    cpu_time = t1_stop - t1_start
    return betas, cpu_time

def linear_regression_gd_stochastic_predict(betas, X_test):
    y_pred = []
    
    #your code below
    # YOUR CODE HERE
    X_test_canonical = np.hstack((np.ones(len(X_test)).reshape(len(X_test),1),X_test))
    beta_reshaped = betas.reshape(X_test_canonical.shape[1],1)
    y_pred = np.dot(X_test_canonical, beta_reshaped)
    return y_pred
    


```

<!-- grade_id: cell-task-12-ans-2, points: 0, grade: false, solution: true -->
```python
betas_stochastic = []
cpu_time_stochastic = 0
#your code below
# YOUR CODE HERE
betas_stochastic, cpu_time_stochastic =linear_regression_gd_stochastic_training(X_train_scaled_std,y_train,
alpha=0.01,nEpoch=50,nIteration=100)
print(betas_stochastic,cpu_time_stochastic)
```

Output:
```
[[ 4.97961092]
 [-0.57493023]
 [-0.16101633]
 [ 0.34625255]
 [ 0.24964117]
 [ 0.45064029]
 [ 0.50582173]
 [ 0.09562893]
 [ 0.04002016]
 [-0.21979051]
 [ 0.56013062]
 [ 0.43619922]
 [ 0.02672886]
 [-0.59437826]
 [ 0.53134123]
 [-0.13312998]
 [-0.61233011]
 [-0.17782757]
 [ 0.15519426]
 [ 0.54494622]
 [ 1.99668838]] 0.08049354099999562
```

<!-- grade_id: cell-task-12-ans-3, points: 0, grade: false, solution: true -->
```python
y_pred = []
#your code below
# YOUR CODE HERE
y_pred =linear_regression_gd_stochastic_predict(betas_stochastic,X_test_scaled_std)

```

<!-- grade_id: cell-task-12-ans-4, points: 0, grade: false, solution: true -->
```python
rmsle_stochastic = 0
#your code below
# YOUR CODE HERE
rmsle_stochastic = RMSLE(y_test,y_pred.flatten())
print(rmsle_stochastic)
```

Output:
```
0.5352490494089946
/var/folders/8w/fqk30xsj2tg4p5hwsm8v1qsc0000gn/T/ipykernel_71457/3094196268.py:31: RuntimeWarning: invalid value encountered in log1p
  rmsle_val = np.sqrt(np.mean(np.square(np.log1p(y_test)-np.log1p(y_pred))))
```

<!-- grade_id: cell-task-12-test, points: 10, grade: true, solution: false -->
```python

```

## Task 13: 
Given the (X_train_scaled_std, y_train) pairs denoting input matrix and output vector respectively: complete the three function definitions and demonstrate the functionalities of each by calling them with appropriate arguments as instructed below:
* **linear_regression_gd_minibatch_training** : It fits a linear regression model using the minibatch gradient descent algorithm to obtain the coefficients, beta's, as a numpy array of m+1 values, where *m* is the number of variables kept in X_train (the first argument to the function). Please use the alpha (i.e, the learning rate), nEpoch (number of epochs), nIteration (number of iterations), and batch_size parameters in your implementation of the gradient descent algorithm. Please measure the cpu_time needed during the training step. cpu_time is not equal to the wall_time. So, use time.perf_counter() for an accurate measurement. Documentation on this function can be found here: https://docs.python.org/3/library/time.html . Finally, the function returns betas (i.e., the m+1 beta values) and the  cpu_time.
* **linear_regression_gd_minibatch_predict**: It takes a list of m+1 beta values (i.e., betas returned from the corresponding training function, and X_test (containing test samples each having *m* input variables). Now, using the provided beta values, predict each of the test samples provided, and let's name your prediction "y_pred". Return y_pred from the function. 

* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES THE LINEAR REGRESSION.
* Now, call linear_regression_gd_minibatch_training() function providing X_train_scaled_std, y_train obtained from Task 9, and alpha=0.01,nEpoch=50, nIteration=1000,batch_size=32, and save the returned results as betas_minibatch,cpu_time_minibatch.
* Print betas_minibatch, cpu_time_minibatch
* Call linear_regression_gd_minibatch_predict() function providing betas_minibatch,X_test_scaled_std obtained in Task 9. Save the returned result as y_pred.
* Call RMSLE() function providing y_test and y_pred. Save returned result as rmsle_minibatch.
* Print rmsle_minibatch.

<!-- grade_id: cell-task-13-ans-1, points: 0, grade: false, solution: true -->
```python
def linear_regression_gd_minibatch_training(X_train,y_train,alpha,nEpoch, nIteration, batch_size):
    random.seed(554433)
    cpu_time = 0
    epsilon = 1e-6
    betas = []
    ## your code below
    # YOUR CODE HERE
    y = np.array(y_train).reshape((len(y_train),1))
    X = np.array(X_train)
    if all(X[:,0]==1) == False:
        X = np.hstack((np.ones(len(X)).reshape(len(X), 1), X))
    betas = np.random.rand(X.shape[1]).reshape((X.shape[1],1))
    convergence_reached = False
    betas_prev = betas
    t1_start = perf_counter()
    counter = 0
    for epoch in np.arange(nEpoch):
        np.random.shuffle(X)
        mini_batches = []
        data = np.hstack((X, y))
        np.random.shuffle(data)
        n_minibatches = data.shape[0] // batch_size
        for i in range(n_minibatches + 1):
            mini_batch = data[i * batch_size:(i + 1) * batch_size, :]
            X_mini = mini_batch[:, :-1]
            Y_mini = mini_batch[:, -1].reshape((-1, 1))
            mini_batches.append((X_mini, Y_mini))
        if data.shape[0] % batch_size != 0:
            mini_batch = data[i * batch_size:data.shape[0]]
            X_mini = mini_batch[:, :-1]
            Y_mini = mini_batch[:, -1].reshape((-1, 1))
            mini_batches.append((X_mini, Y_mini))
        for mini_batch in mini_batches:
            X_mini, y_mini = mini_batch
            error = np.dot(X_mini, betas) - y_mini
            grad = (1 / len(X_mini)) * np.dot(X_mini.T, error)
            betas = betas - alpha * grad
            if counter > nIteration:
                convergence_reached = True
            elif np.sum(np.abs(betas-betas_prev)) < epsilon:
                convergence_reached = True
            betas_prev = betas
            counter = counter + 1
            if convergence_reached:
                break
        if convergence_reached:
            break
    t1_stop = perf_counter()
    cpu_time = t1_stop - t1_start
    
    return betas,cpu_time

def linear_regression_gd_minibatch_predict(betas,X_test):
    #your solution below
    y_pred = []
    
    # YOUR CODE HERE
    X_test_canonical = np.hstack((np.ones(len(X_test)).reshape(len(X_test),1),X_test))
    betas_reshaped = betas.reshape(X_test_canonical.shape[1],1)
    y_pred = np.dot(X_test_canonical, betas_reshaped)
    return y_pred

    




```

<!-- grade_id: cell-task-13-ans-2, points: 0, grade: false, solution: true -->
```python
betas_minibatch = []
cpu_time_minibatch = 0
#your task below

# YOUR CODE HERE
betas_minibatch,cpu_time_minibatch =linear_regression_gd_minibatch_training(X_train_scaled_std,y_train,alpha=0.001,nEpoch=50,nIteration=1000,batch_size=32)
print(betas_minibatch,cpu_time_minibatch)
```

Output:
```
[[ 4.74002510e+00]
 [ 1.66308283e-01]
 [ 1.24785574e-01]
 [ 1.00500393e-01]
 [ 2.72755137e-01]
 [ 5.36125058e-01]
 [ 1.12396186e-01]
 [ 6.76591454e-02]
 [ 3.86744121e-01]
 [-2.95938780e-01]
 [ 2.57586420e-01]
 [ 3.74715698e-03]
 [ 1.40259163e-01]
 [ 2.79271008e-01]
 [ 2.43744419e-01]
 [ 2.20993131e-01]
 [ 2.00171151e-01]
 [ 1.87510042e-02]
 [ 2.04572476e-01]
 [ 1.60486625e-02]
 [ 4.89715410e-02]] 0.17612120799999786
```

<!-- grade_id: cell-task-13-ans-3, points: 0, grade: false, solution: true -->
```python
y_pred = []
#your ans below
# YOUR CODE HERE
y_pred =linear_regression_gd_minibatch_predict(betas_minibatch,X_test_scaled_std)
```

<!-- grade_id: cell-task-13-ans-4, points: 0, grade: false, solution: true -->
```python
rmsle_minibatch = 0
#your ans below
# YOUR CODE HERE
rmsle_minibatch = RMSLE(y_test,y_pred.flatten())
print(len(y_test))
print(len(y_pred))
print(rmsle_minibatch)
```

Output:
```
25350
25350
0.42784842890056574
```

<!-- grade_id: cell-task-13-test, points: 10, grade: true, solution: false -->
```python

```

## Task 14:
Given the 4 sets of results from the 4 experiments (from Tasks 10, 11, 12, 13) with closed form solution, batch gradient descent, stochastic gradient descent and mini-batch gradient descent, assign a string from the set {"closed-form", "batch-GD", "stochastic-GD", "minibatch-GD"} that demonstrated the best predictive performance in terms of RMSE to a variable named `best_name`.


<!-- grade_id: cell-task-14-ans, points: 0, grade: false, solution: true -->
```python
names = ["closed-form", "batch-GD", "stochastic-GD", "minibatch-GD"]
rmsle_vals = [rmsle_closed_form, rmsle_batch, rmsle_stochastic, rmsle_minibatch]

best_name = 'xxxx'
# YOUR CODE HERE
print(min(rmsle_vals))
least = min(rmsle_vals)
for i in range(len(rmsle_vals)):
    if rmsle_vals[i] == least:
        best_name = names[i]
print(best_name)
```

Output:
```
0.13907911709676382
closed-form
```

<!-- grade_id: cell-task-14-test, points: 10, grade: true, solution: false -->
```python

```

## Task 15: 
Given the 4 sets of results from the 4 experiments (from Tasks 10, 11, 12, 13) with closed form solution, batch gradient descent, stochastic gradient descent and mini-batch gradient descent, assign a string from the set {"closed-form", "batch-GD", "stochastic-GD", "minibatch-GD"} that demonstrated the least training cpu time to a variable called `best_name`.


<!-- grade_id: cell-task-15-ans, points: 0, grade: false, solution: true -->
```python
names = ["closed-form", "batch-GD", "stochastic-GD", "minibatch-GD"]
cpu_times = [cpu_time_closed_form, cpu_time_batch, cpu_time_stochastic, cpu_time_minibatch]
best_name = 'xxxx'

# YOUR CODE HERE
print(min(cpu_times))
least = min(cpu_times)
for i in range(len(cpu_times)):
    if cpu_times[i] == least:
        best_name = names[i]
print(best_name)



```

Output:
```
0.008567458999991118
closed-form
```

<!-- grade_id: cell-task-15-test, points: 10, grade: true, solution: false -->
```python

```

## Task 16: 
Given the (X_train_scaled_std, y_train) pairs denoting input matrix and output vector respectively, 
* call your implementation of Task 13: minibatch gradient descent based linear regression for each of these learning rates: {0.001, 0.01, 0.05}, batch sizes: {32, 64, 128, 256}
    * Please use the nIteration (number of iterations)=100, nEpoch (number of epoch)=50.

* For each of the linear regression model, using the computed beta values, predict the test samples provided in the "X_test_scaled_std" argument, and let's name your prediction "y_pred".
* Compute Root Mean Squared Logarithmic Error (RMSLE) of your prediction using the RMSLE() function you defined in Task 10.
* Finally, assign the learning rate that shows the best test performance to a variable called `best_alpha`, and also assign as a pandas dataframe named `summary` with 3 columns: {learning_rate, batch_size, test_RMSLE} containing RMSLE's of the 12 linear regression models, i.e., Cartesian product between the alphas and batch_sizes. Also, assign the best performing batch_size to a variable `best_batch_size`.
* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES THE LINEAR REGRESSION.


<!-- grade_id: cell-task-16-ans, points: 0, grade: false, solution: true -->
```python
alphas = [0.001, 0.01, 0.05]
batch_sizes = [32, 64, 128, 256]
summary = pd.DataFrame(columns=['learning_rate', 'batch_size', 'test_RMSLE'])
nIteration=1000
nEpoch=50

best_alpha = -1
best_beta = []
best_batch_size = -1
a = []
b = []
#your code below
# YOUR CODE HERE
for batch_size in batch_sizes:
    for alpha in alphas:
        betas,cpu_time = linear_regression_gd_minibatch_training(X_train_scaled_std,y_train,alpha,nEpoch, nIteration, batch_size)
        y_pred =linear_regression_gd_minibatch_predict(betas,X_train_scaled_std)
        y_pr = []
        for i in range(len(y_pred)):
            y_pr.append(y_pred[i][0])
            if i == len(y_test) - 1:
                break
        rmsle = RMSLE(y_test,y_pr)
        summary = summary.append(pd.Series([alpha, batch_size, rmsle],index=summary.columns),ignore_index=True)
summary
```

Output:
```
/var/folders/8w/fqk30xsj2tg4p5hwsm8v1qsc0000gn/T/ipykernel_71457/2805722431.py:24: FutureWarning: The frame.append method is deprecated and will be removed from pandas in a future version. Use pandas.concat instead.
  summary = summary.append(pd.Series([alpha, batch_size, rmsle],index=summary.columns),ignore_index=True)
    learning_rate  batch_size  test_RMSLE
0           0.001        32.0    0.413983
1           0.010        32.0    0.193673
2           0.050        32.0    0.195180
3           0.001        64.0    0.402243
4           0.010        64.0    0.193472
5           0.050        64.0    0.194643
6           0.001       128.0    0.421847
7           0.010       128.0    0.193433
8           0.050       128.0    0.193365
9           0.001       256.0    0.418523
10          0.010       256.0    0.193339
11          0.050       256.0    0.193319
```

```python

```

<!-- grade_id: cell-task-16-test, points: 10, grade: true, solution: false -->
```python

```

## Task 17:
* Utilizing the best trained linear regression model (so far), predict the target for each of the samples in the judge_dataset.
    * I believe you will not forget to do the following before call in the prediction algorithm:
        - Save the ID values of the judge dataset into ID_judge and drop it from the judge dataframe.
        - Perform onehot encoding using the same encoder you used to encode X_test (Task 4). 
        - keep only the same top 20 variables as you did in Task 7. 
        - scale the input variables based on the same metrics you used to scale the training dataset (Task 9). 
    * Now, call the prediction function of that model to obtain y_pred.
    * Prepare a pandas dataframe called `result` having columns: {ID, BWEIGHT}, where ID will the ID of the judge sample, and BWEIGHT is the corresponding y_pred value from your model prediction.
    * Save the dataframe as `my_submission.csv`.
* PLEASE DO NOT USE ANY LIBRARY FUNCTION THAT DOES THE LINEAR REGRESSION.


```python
X_j = judge_dataset
```

<!-- grade_id: cell-task-17-ans-1, points: 0, grade: false, solution: true -->
```python
#your task below
#i) store the id in a separate variable and drop it from judge dataset
#ii) use onehot encoder you obtained to encode the judge set
#iii) keep only the top-20 variables in the judge set
#iv) scale using the same scalar you obtained before, and save it to judge_scaled_std variable.
judge_scaled_std = []

# YOUR CODE HERE

ID_judge = X_j[['ID']]
X_j = X_j.drop(['ID'], axis=1)



X_j_ohe = [] 

X_j_ohe, enc_list, categorical_features= lets_do_one_hot_encoding(X_j, categorical_features, True, enc_list)



```

```python
X_j_ohe.columns
```

Output:
```
Index(['SEX', 'MARITAL', 'FAGE', 'GAINED', 'VISITS', 'MAGE', 'FEDUC', 'MEDUC',
       'TOTALP', 'BDEAD', 'TERMS', 'LOUTCOME', 'WEEKS', 'RACEMOM', 'RACEDAD',
       'CIGNUM', 'DRINKNUM', 'ANEMIA', 'CARDIAC', 'ACLUNG', 'DIABETES',
       'HERPES', 'HYDRAM', 'HEMOGLOB', 'HYPERCH', 'HYPERPR', 'ECLAMP',
       'CERVIX', 'PINFANT', 'PRETERM', 'RENAL', 'RHSEN', 'UTERINE',
       'HISPMOM_M', 'HISPMOM_N', 'HISPMOM_O', 'HISPMOM_P', 'HISPMOM_S',
       'HISPMOM_U', 'HISPDAD_M', 'HISPDAD_N', 'HISPDAD_O', 'HISPDAD_P',
       'HISPDAD_S', 'HISPDAD_U'],
      dtype='object')
```

```python
X_j_t20 = X_j_ohe[top20_df.iloc[:,0]]
```

```python
X_j_scaled_std = scaler.fit_transform(X_j_t20) 
len(X_j_scaled_std)
```

Output:
```
2000
```

```python
y_j = y_train.iloc[:2000]
```

```python
y_j
```

Output:
```
16961    9.0000
28131    9.0625
35114    7.1250
85941    6.8125
26819    7.1875
          ...  
57997    7.5625
54695    7.3125
43162    8.1875
28752    8.0625
60913    8.3125
Name: BWEIGHT, Length: 2000, dtype: float64
```

```python
betas_closed_form = []
cpu_time_closed_form = 0
#Your task here, call appropriate function to get these two
# YOUR CODE HERE
i=0
re = []
betas_closed_form, cpu_time_closed_form = linear_regression_closed_form_training(X_j_scaled_std, y_j)
betas_closed_form, cpu_time_closed_form
```

Output:
```
(array([ 7.22837500e+00, -2.44300836e-03, -3.46937583e-02,  1.63984406e-02,
        -1.10778140e-02, -2.05528023e-02,  2.84372940e-02,  2.56233973e-02,
         2.48665488e-02, -4.27607629e-02, -7.80343360e-04, -3.36039738e-02,
        -8.11551877e-02,  4.13151848e-02, -9.69226255e-03, -1.68048278e-02,
         1.06386030e-01, -4.87105520e-04,  3.54333621e-02,  3.69630690e-02,
         4.14378905e-02]),
 0.0009757079999985763)
```

```python
y_pred = []
#Your task here. call appropriate function to get get y prediction for the supplied test dataset scaled
# YOUR CODE HERE
y_pred = linear_regression_closed_form_predict(betas_closed_form,X_j_scaled_std)
```

```python
y_pred
```

Output:
```
array([[7.10568381],
       [7.18689266],
       [7.36203843],
       ...,
       [7.23738895],
       [7.20778647],
       [7.24319151]])
```

<!-- grade_id: cell-task-17-ans-2, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-task-17-ans-3, points: 0, grade: false, solution: true -->
```python
#your task below is to prepare a dataframe `result` with columns: {ID,BWEIGHT} as defined in task-17
result = pd.DataFrame()
# YOUR CODE HERE
result = pd.DataFrame(ID_judge)
result['BWEIGHT'] = y_pred
result.to_csv('my_submission.csv',index=False)
result.head()
```

Output:
```
ID   BWEIGHT
0   1  7.105684
1   2  7.186893
2   3  7.362038
3   4  7.260398
4   5  7.141782
```

<!-- grade_id: cell-task-17-test, points: 10, grade: true, solution: false -->
```python

```

## Task 18: 
* [Submit my_submission.csv at Kaggle](https://www.kaggle.com/competitions/fall2022-birth-weight-prediction/)
* [**Very important**] Please look back from tasks 1-17 to see either do hyperparameter tuning, data analysis, model selection to improve your initial submission.
* Please add details about your best submission by assigning appropriate values from your latest Kaggle submission:
    * name -- your kaggle userid/handle
    * submission_date -- latest date/timestamp of your submission
    * score -- the score you got at kaggle for that submission
    * entries -- number of entries you made/submitted in Kaggle for this task
    * tell_about_the_model -- a string explaining your choice of model, parameters, data analysis, etc.

<!-- grade_id: cell-task-18-ans, points: 30, grade: true, solution: true -->
```python
name = 'Vyanjana Kishaniya'
submission_date = '4th October, 2022'
score = 0.22294
entries = 1
tell_about_the_model = 'My choice of model is Closed-form, as it has least CPU_time and least RMSLE values according to task 14 and 15. Also it provides the best solution and finding a closed-form solution is significantly faster than optimizing using an iterative optimization algorithm like gradient descent.'

#your task is to re-assign proper values based on the task requirement below.
# YOUR CODE HERE

```

```python

```